<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name>
      <text>BankAccount mit Vererbung implementieren (Python)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><strong>Aufgabe</strong>: Implementieren Sie eine Python-Klasse <code>BankAccount</code> als <strong>Unterklasse</strong> einer bereits vorhandenen Elternklasse.</p>

<p><strong>Wichtig:</strong> Die <em>Elternklasse</em> <code>AccountBase</code> ist im Hintergrund bereits vorhanden und darf von Ihnen <strong>nicht</strong> implementiert oder verändert werden. Sie implementieren ausschließlich <code>BankAccount</code>.</p>

<h3>Gegeben (Elternklasse im Hintergrund)</h3>
<ul>
  <li><code>class AccountBase(abc.ABC)</code></li>
  <li>Attribute: <code>self._account_holder</code> (String) und <code>self._balance</code> (Float), simuliert als <em>protected</em>.</li>
  <li>Konstruktor: <code>__init__(self, account_holder, initial_balance=0.0)</code></li>
  <li>Getter-Methoden: <code>get_account_holder(self)</code> und <code>get_balance(self)</code></li>
  <li>Abstrakte Methoden: <code>deposit(self, amount)</code> und <code>withdraw(self, amount)</code></li>
</ul>

<h3>Anforderungen an <code>BankAccount</code></h3>
<ul>
  <li><strong>Vererbung:</strong> <code>class BankAccount(AccountBase):</code></li>
  <li><strong>Konstruktor:</strong> Implementieren Sie <code>__init__(self, account_holder, initial_balance=0.0)</code> und rufen Sie darin den Konstruktor der Elternklasse per <code>super().__init__(...)</code> auf.</li>
  <li><strong>Methoden (überschreiben):</strong>
    <ul>
      <li><code>deposit(self, amount)</code> – Einzahlung: Addiert <code>amount</code> zu <code>self._balance</code> (nur wenn amount > 0).</li>
      <li><code>withdraw(self, amount)</code> – Abhebung: Zieht <code>amount</code> von <code>self._balance</code> ab (nur wenn amount > 0 und der Kontostand ausreicht). Gibt in diesem Fall <code>True</code> zurück, ansonsten <code>False</code>.</li>
      <li><code>__str__(self)</code> – Format: <code>"Kontoinhaber: [Name], Saldo: [Betrag]€"</code></li>
    </ul>
  </li>
</ul>

<p><strong>Hinweis:</strong> Nutzen Sie für den String-Zugriff und die Berechnungen direkt die geerbten Attribute <code>self._account_holder</code> und <code>self._balance</code>.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Überprüfen Sie, dass:</p>
<ul>
<li><code>BankAccount</code> korrekt von <code>AccountBase</code> erbt und per <code>super().__init__()</code> initialisiert wird.</li>
<li><code>deposit</code> und <code>withdraw</code> die geerbte Variable <code>self._balance</code> korrekt verändern.</li>
<li><code>withdraw</code> bei zu geringem Saldo <code>False</code> liefert.</li>
<li><code>__str__(self)</code> exakt dem geforderten Format entspricht.</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>10</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0, 0, 10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>20</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload>class BankAccount(AccountBase):

    def __init__(self, account_holder, initial_balance=0.0):
        # TODO: Elternkonstruktor aufrufen
        pass

    def deposit(self, amount):
        # TODO
        pass

    def withdraw(self, amount):
        # TODO
        return False

    def __str__(self):
        # TODO
        return ""
</answerpreload>
    <globalextra><![CDATA[import abc

class AccountBase(abc.ABC):
    def __init__(self, account_holder, initial_balance=0.0):
        self._account_holder = account_holder
        self._balance = float(initial_balance)

    def get_account_holder(self):
        return self._account_holder

    def get_balance(self):
        return self._balance

    @abc.abstractmethod
    def deposit(self, amount):
        pass

    @abc.abstractmethod
    def withdraw(self, amount):
        pass
]]></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[__student_answer__ = """{{ STUDENT_ANSWER | e('py') }}"""

import subprocess
import os

# Create the test runner script
test_script = """import abc

class AccountBase(abc.ABC):
    def __init__(self, account_holder, initial_balance=0.0):
        self._account_holder = account_holder
        self._balance = float(initial_balance)

    def get_account_holder(self):
        return self._account_holder

    def get_balance(self):
        return self._balance

    @abc.abstractmethod
    def deposit(self, amount):
        pass

    @abc.abstractmethod
    def withdraw(self, amount):
        pass

# Student code
""" + __student_answer__ + """

# Test code
{{ TEST.testcode }}
"""

with open("runner.py", "w", encoding="utf-8") as f:
    f.write(test_script)

# Execute the script
try:
    result = subprocess.run(["python3", "runner.py"], capture_output=True, text=True, check=True)
    print(result.stdout, end="")
except subprocess.CalledProcessError as e:
    print(e.stdout, end="")
    print(e.stderr, end="")
]]></template>
    <iscombinatortemplate>0</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[class BankAccount(AccountBase):

    def __init__(self, account_holder, initial_balance=0.0):
        super().__init__(account_holder, initial_balance)

    def deposit(self, amount):
        if amount > 0:
            self._balance += amount

    def withdraw(self, amount):
        if amount > 0 and self._balance >= amount:
            self._balance -= amount
            return True
        return False

    def __str__(self):
        return f"Kontoinhaber: {self._account_holder}, Saldo: {self._balance}€"
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language>python3</language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs>5</cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>twig</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.5" >
        <testcode>
          <text><![CDATA[acc1 = BankAccount("Alice")
print(acc1.get_balance())
print(acc1.get_account_holder())]]></text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text>0.0
Alice</text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>

      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.5" >
        <testcode>
          <text><![CDATA[acc2 = BankAccount("Bob", 500.0)
print(acc2.get_balance())
print(acc2)]]></text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text>500.0
Kontoinhaber: Bob, Saldo: 500.0€</text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>

      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.5" >
        <testcode>
          <text><![CDATA[acc3 = BankAccount("Carol", 1000.0)
acc3.deposit(250.0)
print(acc3.get_balance())]]></text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text>1250.0</text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>

      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.5" >
        <testcode>
          <text><![CDATA[acc4 = BankAccount("Dave", 800.0)
success1 = acc4.withdraw(300.0)
success2 = acc4.withdraw(600.0)
print(success1)
print(success2)
print(acc4.get_balance())]]></text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text>True
False
500.0</text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>

      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="2.0" >
        <testcode>
          <text><![CDATA[acc5 = BankAccount("Eve", 200.0)
acc5.deposit(300.0)
acc5.withdraw(150.0)
acc5.withdraw(100.0)
print(acc5)]]></text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text>Kontoinhaber: Eve, Saldo: 250.0€</text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>

      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="2.0" >
        <testcode>
          <text><![CDATA[acc6 = BankAccount("Frank")
acc6.deposit(1000.0)
result = acc6.withdraw(1200.0)
print(result)
print(acc6.get_balance())]]></text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text>False
1000.0</text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>
    </testcases>
    <tags>
      <tag><text>OOP</text></tag>
      <tag><text>Vererbung</text></tag>
      <tag><text>Python</text></tag>
      <tag><text>CodeRunner</text></tag>
    </tags>
  </question>
</quiz>
